---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Intro lecture 1

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

Print a linePrint a linePrint a linePrint a linePrint a linePrint a linePrint a linePrint a linePrint a linePrint a line

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
print("Something else")
```

`@solution`
```{python}
print("hello world")
```

`@sct`
```{python}

```
